<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/classic-datasets/Iris.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


# Kata: Iris Dataset

| Learning type | Activity type | Objective |
| - | - | - |
| Supervised | Multiclass classification | Identify a flower's class |

## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Complete the cells beginning with `# YOUR CODE HERE` and run the subsequent cells to check your code.

## About the dataset

[Iris](https://archive.ics.uci.edu/ml/datasets/iris) is a well-known multiclass dataset. It contains 3 classes of flowers with 50 examples each. There are a total of 4 features for each flower.

![](images/Iris-versicolor-21_1.jpg)

## Package setup

```python
# Import needed packages
# You may add or remove packages should you need them
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
from sklearn.datasets import load_iris
from keras.models import Sequential
from keras.layers import Dense
from keras.optimizers import *
from keras.utils import to_categorical

# Display plots inline and change plot resolution to retina
%matplotlib inline
%config InlineBackend.figure_format = 'retina'
# Set Seaborn aesthetic parameters to defaults
sns.set()
```

Output:
```
Using TensorFlow backend.
```

## Utility functions

```python
def plot_loss_acc(history):
    """Plot training and (optionally) validation loss and accuracy"""

    loss = history.history['loss']
    epochs = range(1, len(loss) + 1)

    plt.figure(figsize=(10, 10))

    plt.subplot(2, 1, 1)
    plt.plot(epochs, loss, '.--', label='Training loss')
    final_loss = loss[-1]
    title = 'Training loss: {:.4f}'.format(final_loss)
    plt.ylabel('Loss')
    if 'val_loss' in history.history:
        val_loss = history.history['val_loss']
        plt.plot(epochs, val_loss, 'o-', label='Validation loss')
        final_val_loss = val_loss[-1]
        title += ', Validation loss: {:.4f}'.format(final_val_loss)
    plt.title(title)
    plt.legend()

    acc = history.history['acc']

    plt.subplot(2, 1, 2)
    plt.plot(epochs, acc, '.--', label='Training acc')
    final_acc = acc[-1]
    title = 'Training accuracy: {:.2f}%'.format(final_acc * 100)
    plt.xlabel('Epochs')
    plt.ylabel('Accuracy')
    if 'val_acc' in history.history:
        val_acc = history.history['val_acc']
        plt.plot(epochs, val_acc, 'o-', label='Validation acc')
        final_val_acc = val_acc[-1]
        title += ', Validation accuracy: {:.2f}%'.format(final_val_acc * 100)
    plt.title(title)
    plt.legend()
```

## Step 1: Loading the data

```python
# Load the Iris dataset included with scikit-learn
dataset = load_iris()

# Put data in a pandas DataFrame
df_iris = pd.DataFrame(dataset.data, columns=dataset.feature_names)
# Add target and class to DataFrame
df_iris['target'] = dataset.target
df_iris['class'] = dataset.target_names[dataset.target]
# Show 10 random samples
df_iris.sample(n=10)
```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  petal width (cm)  \
128                6.4               2.8                5.6               2.1   
123                6.3               2.7                4.9               1.8   
18                 5.7               3.8                1.7               0.3   
105                7.6               3.0                6.6               2.1   
0                  5.1               3.5                1.4               0.2   
141                6.9               3.1                5.1               2.3   
11                 4.8               3.4                1.6               0.2   
14                 5.8               4.0                1.2               0.2   
147                6.5               3.0                5.2               2.0   
129                7.2               3.0                5.8               1.6   

     target      class  
128       2  virginica  
123       2  virginica  
18        0     setosa  
105       2  
```

### Question

Store training input data in a variable named `x_train` and training targets in a variable named `y_train`.

<!-- grade_id: cell-68cb077281b12fe7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from sklearn.model_selection import train_test_split
y=df_iris.target
df_iris.drop(['target'], axis=1, inplace=True)
x=df_iris
```

```python
x_train, x_valid, y_train, y_valid = train_test_split(x, y,train_size=0.8, test_size=0.2,random_state=0)
```

<!-- grade_id: cell-c35ec2e0b638d9ec, points: 1, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. y_train: {y_train.shape}')
print(f'Labels: {y_train}')
assert x_train.shape == (150,4)
assert y_train.shape == (150,)
```

Output:
```
x_train: (120, 5). y_train: (120,)
Labels: 137    2
84     1
27     0
127    2
132    2
59     1
18     0
83     1
61     1
92     1
112    2
2      0
141    2
43     0
10     0
60     1
116    2
144    2
119    2
108    2
69     1
135    2
56     1
80     1
123    2
133    2
106    2
146    2
50     1
147    2
      ..
14     0
122    2
19     0
29     0
130    2
49     0
136    2
99     1
82     1
79     1
115    2
145    2
72     1
77     1
25     0
81     1
140    2
142    2
39     0
58     1
88     1
70     1
87     1
36     0
21     0
9      0
103    2
67     1
117    2
47     0
Name: target, Length: 120, dtype: int32
AssertionError                            Traceback (most recent call last)
<ipython-input-20-6d746c4f5b27> in <module>
      1 print(f'x_train: {x_train.shape}. y_train: {y_train.shape}')
      2 print(f'Labels: {y_train}')
----> 3 assert x_train.shape == (150,4)
      4 assert y_train.shape == (150,)

AssertionError:
```

## Step 2: Preparing the data

### Question

Reshape `y_train` to one-hot encode the targets in a (150,3) matrix.

<!-- grade_id: cell-2ee80aeb35907e0c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from keras.utils import to_categorical

tensor = np.matrix([150:3])
y_train = to_categorical (tensor)
```

Output:
```
File "<ipython-input-17-e2029bcf612b>", line 4
    tensor = np.matrix([150:3])
                           ^
SyntaxError: invalid syntax
```

```python
# Show a sample of encoded targets
df_iris_labels = pd.DataFrame(y_train)
df_iris_labels.sample(n=10)
```

<!-- grade_id: cell-be7316c44aaa3e7c, points: 1, grade: true, solution: false -->
```python
print(f'y_train: {y_train.shape}')
assert y_train.shape == (150,3)
assert np.array_equal([1,0,0], y_train[0])
assert np.array_equal([0,1,0], y_train[50])
assert np.array_equal([0,0,1], y_train[100])
```

## Step 3: Training a model

### Question

Train a model on the data to obtain a training accuracy > 93%. Store the training history in a variable named `history`.

Tip: for best results, use the Adam optimizer with a learning rate of 0.1.

```python
model.compile(Adam(lr=0.1), loss='categorical_crossentropy', metrics=['accuracy'])
```

<!-- grade_id: cell-fd837bbf22683090, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Plot training history
plot_loss_acc(history)
```

<!-- grade_id: cell-a93a4cc26df606c8, points: 2, grade: true, solution: false -->
```python
# Retrieve final accuracy
final_acc = history.history['acc'][-1]
# Assert final accuracy
assert final_acc > 0.93
```

```python

```